<a href="https://colab.research.google.com/github/nkouroup-design/class2/blob/main/In_Class_Assignment_3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# In-class assignment 3
## Pandas Basics and Trasforming Data with Pandas

This notebook contains Python problems you need to complete. Write your codes directly here. Save your code and upload this file to Gradescope Note that there may be some hidden tests, so you may not get full points if these tests fail. Do not be misled by the message "q# passed!", which simply means the autograder has successfully graded the question.

## Importing a csv file and gathering information about the dataset
- Use the `read_csv()`, `info()`, `describe` to obtain more information about the dataset.  
- The header and index_col arguments are one of the frequently used options.  Read and understand what they are for.
- For Jupyter Notebook, press SHIFT+TAB for a quick documentation access.  For Google Colab, you just need to move you mouse over the method name, e.g., `read_csv()`.

```python
import numpy as np
import pandas as pd

from google.colab import files
uploaded = files.upload()
```

Output:
```
<IPython.core.display.HTML object>
Saving Basic_Stats.csv to Basic_Stats.csv
```

```python
nfl = pd.read_csv("Basic_Stats.csv", header=0,
                  dtype={"Age":'Int64', "Height (inches)":'Int64', "Weight (lbs)":'Int64'}, skipinitialspace=True)
# Remove the rows if Current Team is NaN
nfl = nfl.dropna(subset=['Current Team'])
# Change the dtype of Birthday to ISO 8601 datetime format
display(nfl.head())
nfl['Birthday'] = pd.to_datetime(nfl['Birthday'], format='%m/%d/%Y', errors='coerce')
nfl = nfl.dropna(subset=['Birthday'])

nfl.info()

```

Output:
```
Age   Birth Place   Birthday           College Current Status  \
4    25  Detroit , MI  3/31/1992  Central Michigan         Active   
7    22  Navarre , FL  1/31/1995           Clemson         Active   
17   25           NaN  3/11/1992        Pittsburgh         Active   
35   29   Dayton , OH   4/1/1988        Ohio State         Active   
37   26    Rodeo , CA  7/28/1990           Ohio U.         Active   

           Current Team  Experience  Height (inches)     High School  \
4   Pittsburgh Steelers  3rd season               77  Clintondale HS   
7         New York Jets      Rookie               77      Navarre HS   
17  Washington Redskins  1st season               74             NaN   
35    Carolina Panthers  8th season               71    Northmont HS   
37      Oakland Raiders  4th season               72  De La Salle HS   

           High School Location              Name  Number  \
4   Clinton Twp.,Macomb Co., MI      Walton, L.T.    96.0   
7                            FL   
```

```python
nfl.describe(include='all').T # Transposed output will be displayed.
```

## Mean and Median

**Question 1 (1pt):**

a. Calculate the mean age of the players and store the result (float) in a variable named `mean_age`.

b. Determine the mean Weight (lbs) of the players and store the result (float) in a variable named `mean_weight`.

```python
# Write your coding answer here
mean_age = nfl['Age'].mean()

mean_weight = nfl['Weight (lbs)'].mean()


...
```

Output:
```
Ellipsis
```

```python
#print(mean_age, mean_weight)
```

## Summarizing

**Question 2 (1.5 pt):** Square brackets provide a simple way to select columns from a DataFrame.

- Identify the earliest birthday (i.e., ***birthday of the oldest player***) present in the dataset and store the result in a variable named `earliest_birthday`.
- Determine the most frequently occurring college in the College column and store the result (String) in a variable named `most_common_college`.
- Summarize the height and weight of the players by using the aggretation functions, `mean`, `std`, and `median`, in that order to display the mean, standard deviation, and median of players's height and weight. The summary should be a DataFrame with two columns, 'Height (inches)' and 'Weight (lbs)', and three rows for mean, std, and median.


```python
# Write your answers here.

earliest_birthday = nfl['Birthday'].min()

most_common_college = nfl['College'].mode().iloc[0]

summary_stats = nfl[['Height (inches)', 'Weight (lbs)']].agg(['mean','std','median'])

...
```

```python
#display(earliest_birthday, most_common_college, summary_stats)
#display(type(earliest_birthday), type(most_common_college), type(summary_stats))
```

## Cumulative Statistics

**Question 3 (1 pt):**
- Calculate the cumulative sum of the players' Weight (lbs) and store the result (a Series object) in a variable named cumulative_weight.
- Determine the cumulative count of non-null entries in the Birth Place column and store the result(a Series object) in a variable named cumulative_birth_place. Hint: Use `.isna()` to detect null entries.  Note that `True` is considered 1 and `False` is considered 0 when arithmetic operations are applied.

```python
# Write your answers here.
cumulative_weight = nfl['Weight (lbs)'].cumsum()

cumulative_birth_place = (nfl['Birth Place'].isna()).cumsum()

...
```

Output:
```
Ellipsis
```

```python
#display(cumulative_weight, cumulative_birth_place)
#display(type(cumulative_weight), type(cumulative_birth_place))
```

## Dropping duplicates & setting and removing index

**Question 4 (1pt):**
- Drop duplicate entries based on the `College` and `High School` columns to create a pandas.DataFrame object showing all unique combinations of `College` and `High School` players attended. Store the result in a variable named `college_hs`.
- Reset the index of `college_hs`, dropping its contents (existing index).

```python
# Write your answers here.
college_hs = nfl[['College','High School']].drop_duplicates().reset_index(drop=True)

...
```

Output:
```
Ellipsis
```

```python
#display(college_hs)
#display(type(college_hs))
```

## Counting Categorical Variables

**Question 5 (1pts):**
- Count the number of occurrences of each unique value in the Position column, sorting the counts in descending order. Store the result in a variable named `position_count`.
- Obtain the proportion of different status (`Current Status`) in nfl, sorting the proportions in descending order. Store the result in a variable named `status_proportion`.
Note that the type of these count and proportion are pandas.Series.

```python
# Write your answers here.
position_count = nfl['Position'].value_counts()

status_proportion = nfl['Current Status'].value_counts(normalize=True)
...
```

Output:
```
Ellipsis
```

```python
#print(type(position_count), type(status_proportion))
#print(position_count)
#print(status_proportion)
```

## Setting multi-level indexes

**Question 6 (1pts):**
- Set the index of nfl to the "Current Team" and "College" columns, and assign this to `nfl_ind`. After then, sort `nfl_ind` by the index values.
- Specify two Current Team/College pairs (***tuples***) to keep: ("Carolina Panthers"/"North Carolina-Charlotte") and ("Buffalo Bills"/"North Carolina State"), assigning to `rows_to_keep`.
- Subset nfl_ind for rows_to_keep using .loc[]. At the same time, use slicing to subset columns from "Experience" to "Weight (lbs)". Store the result in a variable named `my_players`.

```python
# Write your answers here.
nfl_ind = nfl.set_index(['Current Team','College']).sort_index()

rows_to_keep = [("Carolina Panthers","North Carolina-Charlotte"), ("Buffalo Bills","North Carolina State")]

my_players = nfl_ind.loc[rows_to_keep, "Experience":"Weight (lbs)"]
...

```

Output:
```
Ellipsis
```

```python
#display(nfl_ind)
#display(my_players)
#my_players.columns
```

## Slicing time series


**Question 7 (1pt):**
- Set the index of nfl to the "Birthday" column and sort it.  Store the result to `nfl_BD`.
- Use .loc[] to subset nfl_BD for players born in 1996 and 1997. Store the result to `young_players`
- Use .loc[] to subset nfl_BD for players born from Sep 1996 to Aug 1997. Store the result to `class2015`

```python
# Write your answers here.

nfl_BD = nfl.set_index('Birthday').sort_index()

young_players = nfl_BD.loc['1996':'1997']

class2015 = nfl_BD.loc['1996-09':'1997-08']
...
```

Output:
```
Ellipsis
```

```python
#display(nfl_BD)
#display(young_players)
#display(class2015)
```

## Multiple Grouped Summaries

**Question 8 (1pt):**
- Group the DataFrame by Current Team and Position, then calculate the **mean and standard deviation** of `Weight (lbs)` for each group. Store the result in a variable named `grouped_mean_weight`.
- Group the DataFrame by Current Status and College, then find the median Age for each group. Store the result in a variable named `grouped_median_age`.  Make sure the type of the result is pandas.Series for this problem.

```python
# Write your answers here.

grouped_mean_weight = nfl.groupby(['Current Team','Position'])['Weight (lbs)'].agg(['mean','std'])

grouped_median_age = nfl.groupby(['Current Status','College'])['Age'].median()
...
```

Output:
```
Ellipsis
```

```python
#display(grouped_mean_weight)
#display(grouped_median_age)
#display(type(grouped_mean_weight), type(grouped_median_age))
```

## Pivoting on One Variable

**Question 9 (0.5pt):**
- Create a pivot table with Current Team as the index, Position as the columns, and the mean `Height (inches)` as the values. Store the result in a variable named team_position_pivot.


```python
# Write your answers here.

team_position_pivot = nfl.pivot_table(values='Height (inches)', index='Current Team', columns='Position', aggfunc='mean')
...

```

Output:
```
Ellipsis
```

```python
#display(team_position_pivot)
```

## Fill in Missing Values and Sum Values with Pivot Tables

**Question 10 (1 pt):**
- Create a pivot table with College as the index, Current Status as the columns, and the count of Name as the values, filling in missing values with 0. Store the result in a variable named college_status_pivot.
- Sum the values in college_status_pivot to get the total count of players for each Current Status, and store the result (pd.Series object) in a variable named status_totals.

```python
# Write your answers here.

college_status_pivot = nfl.pivot_table(values='Name', index='College', columns='Current Status', aggfunc='count', fill_value=0)

status_totals = college_status_pivot.sum()

...
```

Output:
```
Ellipsis
```

```python
#display(college_status_pivot)
#display(status_totals)
#display(status_totals.loc['Active'])
```

## Pivot nfl by Current Team and Birthyear

**Question 11 (0.5pt):**
- Make a copy of nfl named `nfl_new` using nfl.copy() and then add a Birthyear column to nfl_new, from the year component of the date column.
- Make a pivot table of the count of the Name column, with Current Team as rows, and Birthyear as columns. Assign to `count_by_college_vs_birthyear`, and look at the result.

```python
# Write your answers here.
nfl_new = nfl.copy()
nfl_new['Birthyear'] = nfl_new['Birthday'].dt.year

count_by_college_vs_birthyear = nfl_new.pivot_table(values='Name', index='Current Team', columns='Birthyear', aggfunc='count')
...
```

Output:
```
Ellipsis
```

```python
#display(count_by_college_vs_birthyear)
```

Tests for q1 (points: [0.4, 0.1, 0.4, 0.1])
```
>>> bool(mean_age == nfl.Age.mean())
True
```
```
>>> type(mean_age) == np.float64 or type(mean_age) == np.double
True
```
```
>>> bool(mean_weight == nfl['Weight (lbs)'].mean())
True
```
```
>>> type(mean_weight) == np.float64 or type(mean_weight) == np.double
True
```


Tests for q10 (points: [0.3, 0.01, 0.19])
```
>>> college_status_pivot.equals(nfl.pivot_table(index='College', columns='Current Status', values='Name', aggfunc='count', fill_value=0))
True
```
```
>>> status_totals.equals(college_status_pivot.sum(axis=0))
True
```
```
>>> status_totals.equals(college_status_pivot.sum(axis=0)) or status_totals.equals(college_status_pivot.sum(axis=1))
True
```


Tests for q11 (points: [0.5])
```
>>> count_by_college_vs_birthyear.equals(nfl_new.pivot_table(index='Current Team', columns='Birthyear', values='Name', aggfunc='count'))
True
```


Tests for q2 (points: [0.4, 0.1, 0.4, 0.1, 0.4, 0.1])
```
>>> earliest_birthday == nfl['Birthday'].min()
True
```
```
>>> type(earliest_birthday) == pd.Timestamp
True
```
```
>>> most_common_college == nfl['College'].value_counts().index[0]
True
```
```
>>> type(most_common_college) == str
True
```
```
>>> bool(np.all(summary_stats.values == nfl[['Height (inches)', 'Weight (lbs)']].agg(['mean', 'std', 'median']).values))
True
```
```
>>> type(summary_stats) == pd.DataFrame and bool(np.all(summary_stats.columns.values == np.array(['Height (inches)', 'Weight (lbs)'])))
True
```


Tests for q3 (points: [0.4, 0.1, 0.4, 0.1])
```
>>> bool(np.all(cumulative_weight == nfl['Weight (lbs)'].cumsum()))
True
```
```
>>> type(cumulative_weight) == pd.Series
True
```
```
>>> bool(cumulative_birth_place.iloc[-1] == 2513 and cumulative_birth_place.iloc[5] == 5 and (cumulative_birth_place.iloc[100] == 87))
True
```
```
>>> type(cumulative_birth_place) == pd.Series
True
```


Tests for q4 (points: [0.5, 0.5])
```
>>> college_hs[['College', 'High School']].equals(nfl[['College', 'High School']].drop_duplicates().reset_index(drop=True))
True
```
```
>>> bool(np.all(college_hs.index == pd.RangeIndex(start=0, stop=2576, step=1)))
True
```


Tests for q5 (points: [0.5, 0.5])
```
>>> bool(np.all(position_count.values == nfl['Position'].value_counts(sort=True).values))
True
```
```
>>> bool(np.all(status_proportion.values == nfl['Current Status'].value_counts(sort=True, normalize=True).values))
True
```


Tests for q6 (points: [0.4, 0.1, 0.2, 0.3])
```
>>> nfl_ind.equals(nfl.set_index(['Current Team', 'College']).sort_index())
True
```
```
>>> bool(np.all(my_players.columns.values == pd.Index(['Experience', 'Height (inches)', 'High School', 'High School Location', 'Name', 'Number', 'Player Id', 'Position', 'Weight (lbs)'], dtype='object').values))
True
```
```
>>> rows_to_keep == [('Carolina Panthers', 'North Carolina-Charlotte'), ('Buffalo Bills', 'North Carolina State')]
True
```
```
>>> my_players['Player Id'].iloc[0] == 'austinduke/2558937' and my_players['Player Id'].iloc[1] == 'stevenhauschka/2507374'
True
```


Tests for q7 (points: [0.5, 0.3, 0.2])
```
>>> nfl_BD.equals(nfl.set_index('Birthday').sort_index())
True
```
```
>>> young_players.equals(nfl_BD.loc['1996':'1997'])
True
```
```
>>> class2015.equals(nfl_BD.loc['1996-09':'1997-08'])
True
```


Tests for q8 (points: [0.5, 0.5])
```
>>> grouped_mean_weight.equals(nfl.groupby(['Current Team', 'Position'])['Weight (lbs)'].agg(['mean', 'std']))
True
```
```
>>> grouped_median_age.equals(nfl.groupby(['Current Status', 'College'])['Age'].agg(['median'])['median'])
True
```


Tests for q9 (points: [0.5])
```
>>> team_position_pivot.equals(nfl.pivot_table(index='Current Team', columns='Position', values='Height (inches)', aggfunc='mean'))
True
```
